# Heaps

## Overview

For this assignment you will update the heap data structure implemented in class so that it accepts a `key` function in its initializer, which will allow the contents of the heap to be maintained using an arbitrary priority (as dictated by the key extracted by said function).

See the naive implementation of a priority queue in the [lecture notes](https://braeburn.cs.iit.edu/jupyter/user/lee/notebooks/cs331/source/Lectures/priority-queues.ipynb#2.-Naive-implementation) (and the following examples) for details on how such a `key` function might be used. Your changes to the `Heap` class will need to be more extensive, however, than those found in the naive priority queue, as we are not using a built-in `sort` (which already takes a `key` function). 

Just as with the naive priority queue, the default `key` function should simply sort on the value of the elements themselves — such a default value has already been inserted for you into the `__init__` parameter list. You will, at the very least, need to update the `_heapify` and `add` methods, below, to complete this assignment.

<!-- grade_id: heap, points: 0, grade: false, solution: true -->
```python
class Heap:
    def __init__(self, key=lambda x:x):
        self.data = []
        self.key  = key

    @staticmethod
    def _parent(idx):
        return (idx-1)//2
        
    @staticmethod
    def _left(idx):
        return idx*2+1

    @staticmethod
    def _right(idx):
        return idx*2+2
    
    def _heapify(self, idx=0):
        while True:
            left = Heap._left(idx)
            right = Heap._right(idx)
            max_idx = idx
            if left < len(self.data) and self.key(self.data[left]) > self.key(self.data[idx]):
                max_idx = left
            if right < len(self.data) and self.key(self.data[right]) > self.key(self.data[max_idx]):
                max_idx = right
            if max_idx != idx:
                #swap
                self.data[idx], self.data[max_idx] = self.data[max_idx], self.data[idx]
                idx = max_idx
            else:
                break
                
    def add(self, x):
        self.data.append(x)
        idx = len(self.data) - 1
        pidx = Heap._parent(idx)
        while self.key(self.data[pidx]) < self.key(self.data[idx]) and idx > 0:
            #swap
            self.data[pidx], self.data[idx] = self.data[idx], self.data[pidx]
            idx = pidx
            pidx = Heap._parent(idx)
        
    def max(self):
        return self.data[0]

    def pop_max(self):
        ret = self.data[0]
        self.data[0] = self.data[len(self.data)-1]
        del self.data[len(self.data)-1]
        self._heapify()
        return ret
    
    def __bool__(self):
        return len(self.data) > 0

    def __len__(self):
        return len(self.data)

    def __repr__(self):
        return repr(self.data)
```

<!-- grade_id: heap_test_1, points: 3, grade: true, solution: false -->
```python
from unittest import TestCase
import random

tc = TestCase()
h = Heap()

random.seed(0)
for _ in range(10):
    h.add(random.randrange(100))
tc.assertEqual(h.data, [97, 61, 65, 49, 51, 53, 62, 5, 38, 33])
```

<!-- grade_id: heap_test_2, points: 3, grade: true, solution: false -->
```python
from unittest import TestCase
import random

tc = TestCase()
h = Heap(lambda x:-x)

random.seed(0)
for _ in range(10):
    h.add(random.randrange(100))

tc.assertEqual(h.data, [5, 33, 53, 38, 49, 65, 62, 97, 51, 61])
```

<!-- grade_id: heap_test_3, points: 3, grade: true, solution: false -->
```python
from unittest import TestCase
import random

tc = TestCase()
h = Heap(lambda s:len(s))

h.add('hello')
h.add('hi')
h.add('abracadabra')
h.add('supercalifragilisticexpialidocious')
h.add('0')

tc.assertEqual(h.data,
              ['supercalifragilisticexpialidocious', 'abracadabra', 'hello', 'hi', '0'])
```

<!-- grade_id: heap_test_4, points: 3, grade: true, solution: false -->
```python
from unittest import TestCase
import random

tc = TestCase()
h = Heap()

random.seed(0)
lst = list(range(-1000, 1000))
random.shuffle(lst)

for x in lst:
    h.add(x)

for x in range(999, -1000, -1):
    tc.assertEqual(x, h.pop_max())
```

<!-- grade_id: heap_test_5, points: 3, grade: true, solution: false -->
```python
from unittest import TestCase
import random

tc = TestCase()
h = Heap(key=lambda x:abs(x))

random.seed(0)
lst = list(range(-1000, 1000, 3))
random.shuffle(lst)

for x in lst:
    h.add(x)

for x in reversed(sorted(range(-1000, 1000, 3), key=lambda x:abs(x))):
    tc.assertEqual(x, h.pop_max())
```

```python
def running_medians_naive(iterable):
    values = []
    medians = []
    for i, x in enumerate(iterable):
        values.append(x)
        values.sort()
        if i%2 == 0:
            medians.append(values[i//2])
        else:
            medians.append((values[i//2] + values[i//2+1]) / 2)
    return medians
```

```python
running_medians_naive([3, 1, 9, 25, 12])
```

Output:
```
[3, 2.0, 3, 6.0, 9]
```

```python
running_medians_naive([8, 4, 11, 18])
```

Output:
```
[8, 6.0, 8, 9.5]
```

```python
running_medians_naive([3, 1, 9, 25, 12])
```

Output:
```
[3, 2.0, 3, 6.0, 9]
```

```python
class Min_Heap:
    def __init__(self, key=lambda x:x):
        self.data = []
        self.key  = key

    @staticmethod
    def _parent(idx):
        return (idx-1)//2
        
    @staticmethod
    def _left(idx):
        return idx*2+1

    @staticmethod
    def _right(idx):
        return idx*2+2
    
    def _heapify(self, idx=0):
        while True:
            left = Heap._left(idx)
            right = Heap._right(idx)
            min_idx = idx
            if left < len(self.data) and self.key(self.data[left]) < self.key(self.data[idx]):
                min_idx = left
            else:
                min_idx = idx
            if right < len(self.data) and self.key(self.data[right]) < self.key(self.data[min_idx]):
                min_idx = right
            if min_idx != idx:
                #swap
                self.data[idx], self.data[min_idx] = self.data[min_idx], self.data[idx]
                idx = min_idx
            else:
                break
                
    def add(self, x):
        self.data.append(x)
        idx = len(self.data) - 1
        pidx = Heap._parent(idx)
        while self.key(self.data[pidx]) > self.key(self.data[idx]) and idx > 0:
            #swap
            self.data[pidx], self.data[idx] = self.data[idx], self.data[pidx]
            idx = pidx
            pidx = Heap._parent(idx)
        
    def min(self):
        return self.data[0]

    def pop_min(self):
        ret = self.data[0]
        self.data[0] = self.data[len(self.data)-1]
        del self.data[len(self.data)-1]
        self._heapify()
        return ret
    
    def __bool__(self):
        return len(self.data) > 0

    def __len__(self):
        return len(self.data)

    def __repr__(self):
        return repr(self.data)
```

```python
def running_medians(iterable):
    max_heaps = Heap()
    min_heaps = Min_Heap()
    median = []
    
    for number in iterable:
        # if the heaps are empty, add the number to the max_heap
        if len(max_heaps) == 0 and len(min_heaps) == 0:
            max_heaps.add(number)
            median.append(number)
        elif len(max_heaps) > len(min_heaps):
            if number < max_heaps.max():
                min_heaps.add(max_heaps.pop_max())
                max_heaps.add(number)
            else:
                min_heaps.add(number)
            median.append((max_heaps.max() + min_heaps.min()) / 2)
        elif len(max_heaps) < len(min_heaps):
            if number < max_heaps.max():
                max_heaps.add(number)
            else:
                max_heaps.add(min_heaps.pop_min())
                min_heaps.add(number)
            median.append((max_heaps.max() + min_heaps.min()) / 2)
        elif len(max_heaps) == len(min_heaps):
            if number < max_heaps.max():
                max_heaps.add(number)
                median.append(max_heaps.max())
            else:
                min_heaps.add(number)
                median.append(min_heaps.min())
    return median
```

```python
num = [2,1,3,105,106,2,5]
max_heaps = Heap()
min_heaps = Min_Heap()
for i in num:
    max_heaps.add(i)
    min_heaps.add(i)
print(max_heaps,min_heaps)
```

Output:
```
[106, 105, 5, 1, 3, 2, 2] [1, 2, 2, 105, 106, 3, 5]
```

```python
running_medians(num)
```

Output:
```
[2, 1.5, 2, 2.5, 3, 2.5, 2]
```

```python
# (2 points)

from unittest import TestCase
tc = TestCase()
tc.assertEqual([3, 2.0, 3, 6.0, 9], running_medians([3, 1, 9, 25, 12]))
# (2 points)
```

```python
import random
from unittest import TestCase
tc = TestCase()
vals = [random.randrange(10000) for _ in range(1000)]
tc.assertEqual(running_medians_naive(vals), running_medians(vals))
# (4 points) MUST COMPLETE IN UNDER 10 seconds!
```

Output:
```
/Users/cuilongchang/anaconda/envs/py35/lib/python3.5/unittest/case.py in assertSequenceEqual(self, seq1, seq2, msg, seq_type)
    998         standardMsg = self._truncateMessage(standardMsg, diffMsg)
    999         msg = self._formatMessage(msg, standardMsg)
-> 1000         self.fail(msg)
   1001 
   1002     def _truncateMessage(self, message, diff):

/Users/cuilongchang/anaconda/envs/py35/lib/python3.5/unittest/case.py in fail(self, msg)
    663     def fail(self, msg=None):
    664         """Fail immediately, with the given message."""
--> 665         raise self.failureException(msg)
    666 
    667     def assertFalse(self, expr, msg=None):

AssertionError: Lists differ: [6575[132 chars].5, 5923, 5625.5, 5474, 5401.0, 5474, 5530.0, [6812 chars]62.0] != [6575[132 chars].5, 5328, 5625.5, 5923, 5401.0, 5328, 5530.0, [6812 chars]62.0]

First differing element 20:
5923
5328

Diff is 10029 characters long. Set self.maxDiff to None to see it.
```

```python
import random
from unittest import TestCase
tc = TestCase()
vals = [random.randrange(100000) for _ in range(100001)]
m_mid   = sorted(vals[:50001])[50001//2]
m_final = sorted(vals)[len(vals)//2]
running = running_medians(vals)
tc.assertEqual(m_mid, running[50000])
tc.assertEqual(m_final, running[-1])
```